# Parametric B&eacute;zier spline aerofoil design 

This Jupyter notebook is one of the two notebooks that you have to complete and submit as part of the FEEG1001 Summative Design Assessment. You should first complete the `lapsimulator.ipynb` notebook and then work on this one.

__NOTE: You will have to use the classical Jupyter Notebook interface, rather than the newer Jupyter Lab, for the assessment to work properly.__

Similar to the labs, this notebook has some sections that have been completed for you already, while in some other cells you will have to write some code. There are 5 tasks in total in this notebook, with step-to-step guidance. There are some checks to make sure your code works properly. 

Note that this notebook is a little bit different from other notebooks you are used to in the labs. There are some cells that are read-only and can't be modified. You'll have to write your solution in special cells containing the statements

    # YOUR CODE HERE
    raise NotImplementedError
    # END YOUR CODE HERE
    
replacing these two lines with your own code.

## Problem specification

In this notebook, the task is to design an optimal airfoil geometry for the front or rear wing of a race car competing around the Highfield Street Circuit. From `lapsimulator.ipynb`, we now know we need to minimize drag while producing a target lift (downforce) coefficient times reference area 

$$C_L^* A = 0.52$$

with $A$ the reference area. From your optimisation, you might get a value slightly different from 0.52, but we'll use it here as an reasonable value.

From the Summative Assessment Instructions, the reference area for the front wing is $2000 mm$ (span) $\times$ $500 mm$ (chord), for a total area of $1 m^2$. For the rear wing, the reference area is $1050 mm \times 500 mm = 0.525 m^2$. This means that depending on whether you decided to work on the front or rear wing, the target lift coefficient $C_l^* = 0.52 / A$ for the 2D airfoil can be different. For the front wing, we have $C_l^* = 0.52$, while $C_l^* = 0.99$ for the rear wing.  You can use this notebook to design either airfoil, but you now see that designing a rear wing will be slightly more challenging! In practice, real 3D wings also suffer of finite-wing effects that reduce the overall lift compared to the 2D predictions. However, at this stage of the design using simple (and fast) design tools is sufficient.

To solve this task, you will use concepts similar to those developed in the labs. We will define parametric airfoils and use the $y$ location of one of the control points and the angle of attack as the two design parameters for this design optimisation problem. We will use some of the codes developed in the labs, so you'll have to download from Blackaboard the files `lab_1_tools.py`, `lab_2_tools.py` and `airfoilsim.py` from their Labs folders. As usual, save these files in the same directory with this notebook. Note that you'll have to be connected to the University VPN for the airfoil simulator to work.

Before you proceed, you need to decide if you are going to design a front wing or a rear wing. Edit the cell below and set the variable `cl_star` (used below in the rest of the notebook) to 0.52 (for a front wing) or to 0.99 (for the rear wing).

```python
# THIS CELL CAN BE EDITED

# set cl_star to 0.52 (for a front wing) or to 0.99 (for the rear wing)
cl_star = 0.99
```

```python
# simple sanity check on the value of cl_star. Adjust the cell above until this runs
assert cl_star in (0.52, 0.99)
```

OK, we can now start this notebook with the usual imports.

```python
# Similar to other cells below, this cell cannot be edited.
# If you need to use any other import statements, please put them in editable cells below

# for the maths and arrays
import numpy as np         

# for the plotting
import matplotlib.pyplot as plt

# import simulator
from airfoilsim import airfoilsim

# standard bezier curves
from lab_1_tools import bezier

# plotting function
from lab_2_tools import plot_airfoil
```

### Task 1 - create a parametric B&eacute;zier spline airfoil

In this task you will write a function `parametric_aerofoil_2` that creates a B&eacute;zier spline aerofoil and returns a 2D array `B` of $x,y$ points on the airfoil from the trailing edge, along the lower surface, round the leading edge, and back along the upper surface to the trailing edge. 

This is similar to Lab 2 except now the aerofoil shape should be varied by changing the design parameter $y_{\mathbf{q}_2}$ (defined by a variable `y_q2`), which controls the $y-$location of the upper surface control point `q[2]`. The other coordinates are fixed to default values described below and do not need to be changed. For this task, we'll use standard B&eacute;zier curves, rather than using rational curves, so we do not need to worry about the weights.

I have made a start at this function in the cell below. You should develop code that does what the docstring describes. Complete it and then run the tests in the next cells.

<!-- grade_id: parametric_aerofoil_2, points: 0, grade: false, solution: true -->
```python
# THIS CELL CAN BE EDITED

def parametric_airfoil_2(y_q2, N=71):
    """
    Compute the coordinates of the airfoil defined by the control points 
    
    p = np.array([[1.0, 0.0], [0.5, -0.04], [0.0,  -0.06], [0.00, 0.00]])
    q = np.array([[0.0, 0.0], [0.0,  0.06], [0.12,  y_q2], [1.00, 0.00]])
    
    where the geometry of the lower and upper surfaces are defined by 
    standard Bezier curves. The input argument `y_q2` defines the `y`
    coordinate of the control point `q[2]` on the upper surface.
    
    The function returns the 2D array `B` with the coordinates of the 
    airfoil in the standard format, i.e. from the leading edge down to the
    trailing edge on the lower surface and then back up to the trailing edge 
    on the upper surface, in a clockwise direction. The trailing edge point
    at (1, 0) should be repeated, while the leading edge (0, 0) should only 
    appear once. The function also return the points `p` and `q`.
    
    
    By default, `N=71` points on each surface are utilised, so that 
    `2N-1 = 141` points in total are returned.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return B, p, q
```

Let's try to run this function with $y_{\mathbf{q}_2} = 0.15$ and then plot the result using the `plot_airfoil` function developed in the Labs.

```python
# THIS CELL CAN BE EDITED

# get coordinates and control points
B, p, q = parametric_airfoil_2(0.15)

# plot airfoil and control points
plot_airfoil(B, p, q)
```

In the cell above, try changing the input argument `y_q2` to understand how this parameter changes the geometry and camber of the airfoil. Then run the tests below to check your code.

<!-- grade_id: parametric_aerofoil_2_points_test, points: 6, grade: true, solution: false -->
```python
# check output of the function parametric_aerofoil_2
assert np.allclose(parametric_airfoil_2(0.1)[0][10], [ 0.78717201, -0.01574344], atol=0.005)
```

OK, well done! Now keep going. If you are struggling with this task, have a look at what you did in Lab 2 and then come back here.

### Task 2 - study effect of angle of attack on lift and drag at constant $y_{\mathbf{q}_2}$

In this task, we will examine the effect of the airfoil angle of attack on the lift and drag coefficients by keeping the design parameter $y_{\mathbf{q}_2}$ constant. First, we need to define the Reynolds number at which we are operating. We will consider an average speed on the Highfield Streest circuit of $35 m/s$, a reference chord of $0.5 m$ (same for front or rear wing as per design specification) and a kinematic viscosity of $1.516\times 10^{-5}$ $m^2/s$.

```python
V = 35        # freestream velocity m/s
c = 0.5       # chord m
nu = 1.516e-5 # kinematic viscosity m^2/s
```

In the cell below, complete the definition of the Reynolds number of the wing and then run the test below to check your definition.

<!-- grade_id: reynolds_number_def, points: 0, grade: false, solution: true -->
```python
# THIS CELL CAN BE EDITED

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: reynolds_number_def_2, points: 2, grade: true, solution: false -->
```python
# check definition of Reynolds number
assert np.allclose(Re, 1154353.562005277, atol=0.005)
```

It is convenient to define a function that returns the force coefficients from the parameter $y_{\mathbf{q}_2}$ and the angle of attack. These are our two design parameters. 

In the cell below, I have defined the signature of this function and wrote the docstring. Complete the code and then run the tests below.

<!-- grade_id: force_coeffs, points: 0, grade: false, solution: true -->
```python
# THIS CELL CAN BE EDITED

def force_coefficients(y_q2, alpha, Re):
    """
    Return lift and drag coefficients of the airfoil geometry
    produced by the function `parametric_airfoil_2` with 
    argument `y_q2` at an angle of attack `alpha` and Reynolds number `Re`.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return cl, cd
```

<!-- grade_id: force_coeffs_tests, points: 3, grade: true, solution: false -->
```python
# check force_coefficients function
assert np.allclose(force_coefficients(0.1, 0, 1e6), (0.0855651913, 0.0055777835), atol=0.005)
```

OK, in the next cell you'll have to write code that runs the `force_coefficients` function for 15 equally spaced angles of attack between -6 and 12, at constant $y_{\mathbf{q}_2} = 0.1$ and for the Reynolds number defined above. You should define an array `alphas` for the angles of attack and store the lift and drag coefficients into separate lists or 1D arrays with name `cls` and `cds`, respectively. There are various techniques you could use, e.g. a for loop, a list comprehension, ...

Then run the tests below to check your code. 

<!-- grade_id: sweep, points: 0, grade: false, solution: true -->
```python
# THIS CELL CAN BE EDITED

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: sweep_tests, points: 3, grade: true, solution: false -->
```python
# check parameter sweep function
assert len(alphas) == 15
assert np.allclose(alphas[2], -3.4285714285714284, atol=0.005)
assert np.allclose(   cls[4],  0.0021465351, atol=0.005)
assert np.allclose(   cds[2],  0.0061810222, atol=0.005)
```

OK, in the cell below we plot the results of this analysis. If your code is correct, you should see the lift increase almost linearly with angle of attack, and the drag coefficient having a minimum at around $\alpha = 0$.

```python
# increase `w`idth space between subplots
plt.subplots_adjust(wspace=0.45)

# make first subplot for lift
plt.subplot(121)

plt.plot(alphas, cls)

# add labels and grid
plt.xlabel("alpha")
plt.ylabel("$C_l$")
plt.grid(1)

# add axes
plt.axvline(0, zorder=0, color="gray")
plt.axhline(0, zorder=0, color="gray")


# make second subplot for drag
plt.subplot(122)

plt.plot(alphas, cds)

# add labels and grid
plt.xlabel("alpha")
plt.ylabel("$C_d$")
plt.grid(1)

# add axes
plt.axvline(0, zorder=0, color="gray")
plt.axhline(0, zorder=0, color="gray")
```

OK, well done! Keep going now.

### Task 3 - study effect of angle of attack and $y_{\mathbf{q}_2}$ on lift and drag coefficients

In this task we will explore the parameter space of this design problem, by computing force coefficients for various combinations of co-ordinate $y_{\mathbf{q}_2}$ and angle of attack. We will produce 2D contour plots as in Lab 4.

In the cell below, you will have to develop some code that loops over an array `alphas` with 8 angles of attacks equally spaced between -2 to 8 degrees and an array `y_q2s` with 9 co-ordinates $y_{\mathbf{q}_2}$ from 0 to 0.35 and evaluates the force coefficients for each possible combination. 

You should initialise two 2D arrays, `CLS` and `CDS`, with shape (8, 9) and use the two-nested-for-loops technique to fill them with the lift and drag coefficients, respectively. 

Use the suggested names for these arrays, so that the test below might have a chance to pass.

<!-- grade_id: nested-for-loop, points: 0, grade: false, solution: true -->
```python
# THIS CELL CAN BE EDITED

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: nested-for-loop_tests, points: 5, grade: true, solution: false -->
```python
# checks code in previous cell
assert len(alphas) == 8
assert len(y_q2s)  == 9
assert alphas[2]   == 0.8571428571428572
assert  y_q2s[2]   == 0.0875
assert CLS[2, 5]   ==  0.4144238202
assert CDS[5, 7]   == 0.0109779256
```

OK, now we plot the result, using the `contourf` and `contour` functions we used in Lab 4. No need to add any code here. Just make sure you understand the code and look carefully at the plots.

```python
# create figure
plt.figure(1, figsize=(12, 4))

# adjust space between subplots
plt.subplots_adjust(wspace=0.3)

## make first subplot
plt.subplot(121)

# plot 30 filled contours of the lift coefficient
plt.contourf(y_q2s, alphas, CLS, 30, cmap="Reds")
plt.colorbar(label="$C_l$", ticks=[-0.3, 0, 0.3, 0.6, 0.9, 1.2])

# add black contours at Cl = Cl_star for the front and rear wings
lab = plt.contour(y_q2s, alphas, CLS, [0.52, 0.99], colors="k")
plt.clabel(lab, fmt="$C_l^* = %.2f$")

# add labels
plt.xlabel("$y_{\mathbf{q}_2}$")
plt.ylabel("alpha")


## make second subplot
plt.subplot(122)
plt.contourf(y_q2s, alphas, CDS, 40, cmap="gist_earth_r")
plt.colorbar(label="$C_d$", ticks=[0.006, 0.008, 0.010, 0.012, 0.014, 0.016])

# add black contours at Cl = Cl_star for the front and rear wings
lab = plt.contour(y_q2s, alphas, CLS, [0.52, 0.99], colors="k")
plt.clabel(lab, fmt="$C_l^* = %.2f$")

# add labels
plt.xlabel("$y_{\mathbf{q}_2}$")
plt.ylabel("alpha")
```

OK, so we see that the minimum drag would occur for $y_{\mathbf{q}_2} \approx 0.1$ and `alpha = 0.5`. However, for these parameter values, the lift would not be high enough for our racecar. 

I have also added two contour lines for the target lift coefficient $C_l^*$ (0.52 for the front wing and 0.99 for the rear wing).

We would like then to find a better combination of these two parameters that satifies the constraint $C_l = C_l^*$ and which minimizes drag. Can you already guess what this should be?

### Task 4 - constrained optimisation - defining the modified objective function

To address this task, we are going to use the optimisation tools developed in Lab 4. Formally we want to solve

\begin{equation}
    \begin{aligned}
        &\underset{ y_{\mathbf{q}_2}, \alpha }{\text{min}}&& C_d{(y_{\mathbf{q}_2}, \alpha)}\\
        &\text{subject to} && C_l{(y_{\mathbf{q}_2}, \alpha)} = C_l^* \\
    \end{aligned}
\end{equation}


We will solve this constrained problem by using the `scipy.optimize` function `minimize`, using the penalty technique to enforce the constraint on the lift coefficient, as in Lab 4.


The first step is to define the function `drag_modified` which takes a two-element 1D array or list containing the co-ordinate `y_q2` and the angle of attack and returns the sum of drag coefficient plus a penalisation terms for the violation of the lift constraint. You should use a penalisation weight $\gamma = 10$.

<!-- grade_id: drag_modified, points: 0, grade: false, solution: true -->
```python
# THIS CELL CAN BE EDITED

def drag_modified(params):
    """ 
    For airfoil geometries obtained with the `parametric_airfoil_2` function 
    defined above, return the sum of drag coefficient plus a penalisation term 
    for the violation of the lift constraint `cl = cl_star`, where `cl_star` is 
    a global variable defined at the beginning of the notebook.
    
    The input argument `params` is a two-element array or list containing 
    the co-ordinate `y` for control point q_2 and the angle of attack `alpha`.
    
    A penalisation weight $\gamma = 10$ must be used.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return J
```

OK, run the tests below to check your implementation.

<!-- grade_id: drag_modified_test, points: 2, grade: true, solution: false -->
```python
# checks 
if cl_star == 0.52:
    assert drag_modified([0.0, 0.1]) == 4.2342601502503765
else:
    assert drag_modified([0.0, 0.1]) == 12.555446030010376
```

### Task 5 - constrained optimisation - running the optimizer

In this last task, we will run the `Nelder-Mead` optimizer using the modified drag function defined above, with tolerance `tol=1e-4`. This is one of the most challenging tasks, so it fine if this will take you a bit of time. Check what we did in Lab 4, then come back here. This step will give us an airfoil that can compete well around the Highfield street circuit.

We first import the required module.

```python
from scipy.optimize import minimize
```

We then define a suitable initial guess `params_guess` for the angle of attack and the co-ordinate $y_{\mathbf{q}_2}$, based on the contour plots. This is already close to the optimum, and we'll use the optimiser to refine the value.

```python
if cl_star == 0.52:
    params_guess = [0.20, 3]
else:
    params_guess = [0.22, 6]
```

In the cells below, you will have to write code to call the optimizer with `drag_modified`. You should store the output of the function `minimize` into a variable `out`, for the checks below to work.

<!-- grade_id: optimizer, points: 0, grade: false, solution: true -->
```python
# THIS CELL CAN BE EDITED

# YOUR CODE HERE
raise NotImplementedError()
```

OK, let's check the output first and we then run some tests to check your implementation is correct. Note that the variable `out.x` will contain the optimal co-ordinate $y_{\mathbf{q}_2}$ and the optimal angle of attack.

```python
print(out)
```

<!-- grade_id: optimizer_test, points: 4, grade: true, solution: false -->
```python
# check solution of optimisation problem
assert (out.x[0] >  0.00) and (out.x[0] < 0.35)
assert (out.x[1] > -2.00) and (out.x[1] < 8.00)
```

OK, we now plot the result of the optimisation, with a red star to denote the optimum. We also save a figure with filename `optfoil.png` in the current directory.

```python
# create figure
plt.figure(1, figsize=(12, 4))

# adjust space between subplots
plt.subplots_adjust(wspace=0.3)

## make first subplot
plt.subplot(121)

# plot 30 filled contours of the lift coefficient
plt.contourf(y_q2s, alphas, CLS, 30, cmap="Reds")
plt.colorbar(label="$C_l$", ticks=[-0.3, 0, 0.3, 0.6, 0.9, 1.2])

# add black contours at Cl = Cl_star for the front and rear wings
lab = plt.contour(y_q2s, alphas, CLS, [0.52, 0.99], colors="k")
plt.clabel(lab, fmt="$C_l^* = %.2f$")

# add labels
plt.xlabel("$y_{\mathbf{q}_2}$")
plt.ylabel("alpha")


## make second subplot
plt.subplot(122)
plt.contourf(y_q2s, alphas, CDS, 40, cmap="gist_earth_r")
plt.colorbar(label="$C_d$", ticks=[0.006, 0.008, 0.010, 0.012, 0.014, 0.016])

# add black contours at Cl = Cl_star for the front and rear wings
lab = plt.contour(y_q2s, alphas, CLS, [0.52, 0.99], colors="k")
plt.clabel(lab, fmt="$C_l^* = %.2f$")

# add labels
plt.xlabel("$y_{\mathbf{q}_2}$")
plt.ylabel("alpha")

# plot optimum point with a red star
plt.plot(out.x[0], out.x[1], "r*", ms=20)

# save figure
plt.savefig("optfoil.png", dpi=600)
```

OK. With this notebook you have now designed an optimal airfoil geometry for your racecar wing and found the optimal angle of attack. 

You can now use these results and bring your design to a further stage of the design process.

__NOTE: Restart the kernel and run all (click "Kernel", "Restart & Run All"), to make sure everything is working correctly before submitting this .ipynb file via Blackboard.__